Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = Steven Tey
COLLABORATORS = 
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
list = []
def bruteforce_max_subarray(A):
    
    for i in range (0, len(A)):
        for j in range (i+1, len(A)+1):
            list.append(sum(A[i:j]))
    
    for i in range (0, len(A)):
        for j in range (i+1, len(A)+1):
            if sum(A[i:j]) == max(list):
                return i, j-1, sum(A[i:j])

A = [-2, -3, 4, -1, -2, 1]
print(f"The value of the maximum subarray of A is {bruteforce_max_subarray(A)[2]}.")
```

Output:
```
The value of the maximum subarray of A is 4.
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
Since the number that is added to the subarray is bigger than the value of the maximum subarray, we will only need to run the bruteforce maximum subarray algorithm from the first term of the previous maximum subarray to the last new appended number. 

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False. If the new maximum subarray value is different, that means that the added term must have added all or some of its value to the new max-subarray value, and thus, the new maximum-subarray has to include the newly appended element, 8.

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Inputs:
    - x: a NON-EMPTY list of numbers (e.g., the array B in the first two questions above). 
        * If x has 1 element: returns the value of the element regardless of the value of mx
    - mx: the maximum subarray of x excluding its last element (i.e., compute the 
    maximum subarray of the input array x considering only its first len(x) - 1 elements)
    
    Output: The maximum subarray of the array x.
    
    For example, using the array B in question 2, the result of incremental_max_subarray(B, 4) 
    is 10 (10 = 8 + 1 - 2 -1 + 4). 
    """
    list = []
    if (len(x)) == 1: #if there is only 1 number in the array, that number is the sum of the max subarray
        return x[0]
    else:
        for i in range(1,len(x)+1): #iterating from 1 to the amount of numbers in the array
            if sum(x[-i:])> mx: #counting backwards - if the sum of the last number, last two numbers,..., is > mx 
                list.append(sum(x[-i:])) # add that number to the list
            else:
                list.append(mx) #else, mx is still the sum of the max subarray
    return max(list)  
    
print(incremental_max_subarray([-2, -3, 4, -1, -2, 1, 8], 4))
```

Output:
```
10
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

```python
B = [-2, -3, 4, -1, -2, 1, 8]
print(incremental_max_subarray(B, 4) == 10)
print(incremental_max_subarray(B[:1], 0) == B[0])
```

Output:
```
True
True
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
A = [-2, -3, 4, -1, -2, 1, 8]
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: flaot, the sum of the maximum subarray of A
    """
    max_list = []
    for x in range(0, len(A)): #iterating across the list from the first to the last element
        max_list.append(incremental_max_subarray(A[:x+1],sum(A[:x]))) #applying the 'incremental_max_subarray'
                                                                      #strategy iteratively across A to produce
                                                                      #the maximum subarray of A
    return max(max_list)
        
print(max_subarray(A))      
```

Output:
```
10
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
Based on the readings, we learned that the divide and conquer algorithm takes $nlogn$ time to complete.

On the other hand, the incremental maximum subarray strategy above uses a recurrence relation, which can be denoted as follows:

$T(n) = T(n-1) + n + c$

The first step will take $T(n)$ time to calculate.

Then, we as go one step down, the time taken will be $T(n-1)$ since we are taking away one term.

Therefore, if you keep iterating, you'll get the final sum to be:

$T(n) + T(n-1) + T(n-2) + ... + T(1)$

This means that the total time needed for this strategy is $T(n+(n-1)+(n-2)+...1) = \frac{1}{2}
n(n+1)$ = $\theta(n^2)$, which is way larger than $nlogn$.

Therefore, this strategy is still not more efficient than the divide-and-conquer strategy.