Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = Steven
COLLABORATORS = None
```

---

# CS110 Pre-class Work 1.2



## Question 1
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertionSort(A):
    for j in range(len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
            A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertionSort(A):
    for j in range(len(A)): #scans the whole array A from first to last position
        key = A[j] #the number under scrutiny
        i= j-1 #the number before the number under scrutiny
        while i >= 0 and A[i]>key: #the loop will keep going if the prev. no. is bigger than the no. under scrutiny
            A[i+1] = A[i] #key switches places with the previous number
            i -= 1 #keeps checking for the previous number to see if this condition is met or not
        A[i+1] = key 
    return A
#this while loop is what differs between an insertion sort and a bubble sort, as it allows the algorithm to 
#continuously compare the key to previous numbers until it is situated in the right position (where it is no longer
#bigger than the previous number)
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertionSort([0]) == [0])
assert(insertionSort([-1,1]) == [-1,1])
assert(insertionSort([1,-1]) == [-1,1])
assert(insertionSort([1,6,3,6]) == [1,3,6,6])
```

```python
def insertionSort(A):
    for j in range(len(A)): 
        key = A[j] 
        i= j-1 
        while i >= 0 and A[i]>key: 
            A[i+1] = A[i] 
            i -= 1 
        A[i+1] = key
    return A

print(insertionSort([0]) == [0])
print(insertionSort([-1,1]) == [-1,1])
print(insertionSort([1,-1]) == [-1,1])
print(insertionSort([1,6,3,6]) == [1,3,6,6])
```

Output:
```
True
True
True
True
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
I tested it with the 4 examples that were given in the assert() functions and they all turned out to be true. 

## Question 3 
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertionSort`,  your new Python function, `bubbleSort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubbleSort(A):
    x = len(A)
    for i in range(x):
        for j in range(0, x-i-1):
            if A[j] > A[j+1] :              
                A[j+1], A[j] = A[j], A[j+1]
    return A

#Here, we are using for loops because we know exactly how long the list and how many positions we need to run the
#algorithm through.

print(bubbleSort([0]) == [0])
print(bubbleSort([-1,1]) == [-1,1])
print(bubbleSort([1,-1]) == [-1,1])
print(bubbleSort([1,6,3,6]) == [1,3,6,6])
```

Output:
```
True
True
True
True
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
Like with the insertion sort algorithm, I ran the bubble sort algorithm with the 4 lines from the assert() section earlier and they turned out to be true.

## Question 4
The following pseudocode corresponds to a third sorting algorithm to be covered in Lesson 1.2, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free. `selectionSort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selectionSort(A):
    for i in range(len(A)): 
        smallest_number = i 
        for j in range(i+1, len(A)): 
            if A[smallest_number] > A[j]: 
                smallest_number = j 
      
        A[i], A[smallest_number] = A[smallest_number], A[i]
    
    return A

#Here, we are using for loops because we know exactly how long the list and how many positions we need to run the
#algorithm through.

print(selectionSort([0]) == [0])
print(selectionSort([-1,1]) == [-1,1])
print(selectionSort([1,-1]) == [-1,1])
print(selectionSort([1,6,3,6]) == [1,3,6,6])
```

Output:
```
True
True
True
True
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**