Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = Steven Tey
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.1

## Question 1.

### Question 1a.

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import numpy as np
import random

counter = []

def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Input:
    - A: list of numerical values
    
    Output:
    - A: Sorted list
    """
    while A != sorted(A):
        np.random.shuffle(A)
        counter.append(1)
    
    return A, len(counter)

print(worst_sort([2, 5, 2, 6, 7, 8, 3]))
```

Output:
```
([2, 2, 3, 5, 6, 7, 8], 300)
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b.
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
This is for the one very lucky case where the first shuffle that the algorithm does happens to shuffle the list into perfect order, or even better, the given list itself is already sorted. In this case, the time complexity will just be $O(n)$, since the algorithm will still have to compare individual elements within the list to find out if they are sorted or not (DISCLAIMER: this is not how I did it in the previous cell, because I used the ``sorted()`` function in Python for that). 

Fun fact: The probability of the original input list being in the exact order it's in is $\frac{1}{n!}$.

### Question 1c.

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
The average case complexity for this is $O((n+1)!)$, since that is the expected number of comparisons and shuffles that needs to be done to get a sorted list of numbers. This number grows exponentially as the size of the list increases. 

### Question 1d.

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
Since the probability of this method working decreases exponentially as the list size increases, I'd say that this method is most feasible when the list size is as small as possible. 

Another feasible scenario is if all the elements in the list are the same.

### [Optional] Question 1e.

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
We could make the ``worse_sort()`` algorithm even more inefficient by making it call on itself recursively with smaller and smaller subsets of the beginning of the list to see if they are sorted. First, it sorts the first $n-1$ elements of the list using this sorting algorithm (recursion). Then, it checks to see if the $n$-th element of the sorted copy is greater than the highest element of the first $n-1$ elements. If so, the copy is now sorted, else randomise the order of the elements of the copy and go nack to the previous recursive check step. Lastly, it check to see if the copy is in the same order as the original list.

According to wikipedia, this is an algorithm that was designed not to succeed before the _heat death of the universe_ on any sizable list. 

## Question 2.
Approximate median finder. Given a list and δ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the (50-δ/2)% and (50+δ/2)% percentiles. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a.
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is within (50-delta/2)% and 
    (50+delta/2)% percentiles.
    X% percentile of a list is defined to be the smallest number in the list that is as large as at least
    X% numbers in the list. 
    
    Input:
    - A: list of numerical values
    - num: the number of interest
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - isApproxMed: a boolean value indicating whether num is within the bound
    """
    larger_numbers = 0
    smaller_numbers = 0
    
    for i in range(len(A)):
        if A[i] >= (num - 0.25):
            larger_numbers += 1
        if A[i] <= (num + 0.25):
            smaller_numbers += 1
    if len(A)%2 == 1: 
        if larger_numbers == smaller_numbers:
            return True
        else:
            return False
    elif len(A)%2 == 0:
        if larger_numbers + 1 == smaller_numbers or larger_numbers == smaller_numbers + 1:
            return True
        else:
            return False
```

```python
print(check_approx_med([0,1], 0, .25) == True)
print(check_approx_med([0], 0, .25) == False)
```

Output:
```
True
False
False
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0,1], 0, .25) == True)
assert(check_approx_med([0], 0, .25) == False)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-68-aa1f56e94d4b> in <module>
      1 assert(check_approx_med([0,1], 0, .25) == True)
----> 2 assert(check_approx_med([0], 0, .25) == False)

AssertionError:
```

### Question 2b.
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 50-delta/2% and 50+delta/2% percentiles
    
    Input:
    - A: list of numerical values
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - num: the approximated median, if it is found within 100//delta trials (this is one possible 
    maximum number of retries. While you are encouraged to play around with another limits, the 
    submitted work must use this number.)
    - None: finding failed, if nothing found within 100//delta trials

    Note: the 100//delta is chosen here as it is the expected number of trials for the first successful finding to occur (using geometric distribution)
    """
    for _ in range(int(100//delta)):
        random_index = random.randint(0,len(A)-1)
        if check_approx_med(A, A[random_index], delta) == True:
            return A[random_index]
    raise ValueError("None")
    

print(approx_med_finder([2, 5, 2, 6, 7, 8, 3], 0.25))    
```

Output:
```
5
```

### Question 2c.

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with δ or N (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
The probability is failure depends on the delta, but it should be around $1-\delta$ since the probability of success is $\delta$. 

This probability should scale with an inverse relation to $\delta$, since the bigger the boundary condition, the lower the probability of failure.

### Question 2d.
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
The expected runtime of ``approx_med_finder`` should be faster than the runtime for a regular sort algorithm to find the median of the list. 