Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = Feng Nian Tey (Steven)
COLLABORATORS = None
```

---

# CS110 Pre-class Work 2.2

## Question 1 (Exercise 3.1-3 of Cormen, et al. )
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
This is because the term $O(n^2)$, also known as the “big-O notation", is used to give an upper bound on a function, to within a constant factor. In other words, $O(n^2)$ is the upper bound on the the worst-case running time for a particular function $g(n)$, and not the minimum running time.

## Question 2 (Exercise 3.1-4 of Cormen, et al. )

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
To solve this, we need to recall the definition of the big-O notation -  it is a function whose curve on a graph will always exceed the curve of the runtime after a certain point. In mathematical prose, it means that $f(n) = O(g(n))$ if and only if there are positive constants $c$ and $k$, such that $0 ≤ f(n) ≤ cg(n)$ for all $n ≥ k$. The values of the constants $c$ and k must be fixed for the function f and must not depend on n.

In the first equation, $2^{n+1}=O(2^n)$, we can factorize the LHS of the equation to:

<div style="text-align:center">$2\times2^{n}=O(2^n)$</div>

Here, we can see that the "$2$" that we factorized out is a constant term, and therefore, we can multiply the $2^n$ term on the RHS by any constant and it can either be equal or bigger than the LHS.

In other words, with $c = 2$ and $k = 1$, we have $2\times2^n \geq 2^{n+1}$ for all $n \geq 1$. Therefore , $2^{n+1}=O(2^n)$.

As for the second example, $2^{2n}=O(2^n)$, we can once again factorize the LHS of the equation to:

<div style="text-align:center">$2^{n\times n}=O(2^n)$</div>
<div style="text-align:center">$2^{n}\times2^{n}=O(2^n)$</div>

Here, we can see that the "$c$" term, $2^{n}$ is not a constant, and therefore, even if we multiplied $2^n$ term on the RHS by any constant, it might still be smaller than the LHS. Thus, $2^{2n}\neq O(2^n)$

## Question 3.
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
list = []
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats 
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    for i in range (0, len(A)): #iterates through the list from the first number to the last
        for j in range (i+2, len(A)+1): #same as the previous step, but this iterates from the term to the ith term 
                                        #to the last term in the list
            list.append(sum(A[i:j])) #add the list of sums to a list
    
    for i in range (0, len(A)): #iterate again
        for j in range (i+2, len(A)+1):
            if sum(A[i:j]) == max(list): #find the pair of integers in the list that give the biggest sums
                return i, j-1, sum(A[i:j]) #return the index of said integers, as well as the sum

A = [-2,1,-1,2,-5]
print(bruteforce_max_subarray(A))

# I acknowledge that there might be limitations of this algorithm that I created - for example, if the list was
# [-1, -1, -1, 2, -1, -1], the algorithm will return "None" because it won't be able to find the maximum difference 
# since there are multiple of them. We would need to add a contingency by using an if statement in order to prevent
# anything like this from happening.
```

Output:
```
(1, 3, 2)
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

```python
print(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

Output:
```
True
```

## Question 4. 
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
list = []
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    for i in range (0, len(A)):
        for j in range (i+2, len(A)+1):
            list.append(sum(A[i:j]))
    
    for i in range (0, len(A)):
        for j in range (i+2, len(A)+1):
            if sum(A[i:j]) == max(list):
                return i, j-1, sum(A[i:j])

A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
print(bruteforce_max_subarray(A))
```

Output:
```
(7, 10, 43)
```

## Question 5. Asymptotic notation. 
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |  |
Selection sort |  |
Bubble sort |  | 
Finding maximum subarray |  |

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->
Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |$O(n^2)$ | $\Theta(n^2)$
Selection sort | $O(n^2)$ |$\Theta(n^2)$
Bubble sort |$O(n^2)$  | $\Theta(n^2)$
Finding maximum subarray* |$O(n^2)$/$O(n^3)$ |$\Theta(n^2)$/$\Theta(n^3)$

<div style="text-align:center">* values for Finding maximum subarray might differ based on the amount of nested for loops that are used.</div>

## [Optional] Question 6. 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
list = []
def bruteforce_min_subarray(A):
    """Implements brute-force minimum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the min subarray
    - the end index of the min subarray
    - the value of the minimum subarray
    """
    for i in range (0, len(A)):
        for j in range (i+2, len(A)+1):
            list.append(sum(A[i:j]))
    
    for i in range (0, len(A)):
        for j in range (i+2, len(A)+1):
            if sum(A[i:j]) == min(list):
                return i, j-1, sum(A[i:j])

A = [-2,1,-1,2,-5]
print(bruteforce_min_subarray(A))
```

Output:
```
(0, 4, -5)
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```